Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Jacob Ravenberg"
COLLABORATORS = ""
```

---

# Fourier Transform 2020

```python
import numpy as np


def sinc(x):
    return np.sinc(x / np.pi)
```

Note:  if you need to use the sinc function, please use the one provided above rather the one in the NumPy library.

## Part A

Find the Fourier transform for $f(t)$ which is defined piecewise
$$
   f(t)=
    \begin{cases}
      -t & -2\leq t \leq 0\\
      t & 0 < t \leq 2\\
      0 & \text{otherwise}
    \end{cases}.
$$
Represent your answer as a function `F(omega)`.

<!-- grade_id: cell-6728c774e597c7e2, points: 0, grade: false, solution: true -->
```python
def F(omega):
    if omega == 0:
        return 4.0
    return ((4*np.sin(2*omega))/omega) - (2*(1-np.cos(2*omega)))/(omega**2)
```

<!-- grade_id: cell-27b08a3d599bb3fe, points: 2, grade: true, solution: false -->
```python
assert isinstance(F(0), (float, complex))


```

## Part B

Find the inverse transform $f(t)$ if $|F(\omega)|$ and $\angle F(\omega)$ are given below.  

<img src="fourier_transform2020.png" width="600"/>

Represent your answer as a function `f(t)`.  

<!-- grade_id: cell-767b795193e9fe8e, points: 0, grade: false, solution: true -->
```python
def f(t):
    return (2*np.cos(2*(t-1))*np.sin(t-1))/(np.pi*(t-1))
```

<!-- grade_id: cell-92bcf1ca08f7e2b6, points: 2, grade: true, solution: false -->
```python
assert isinstance(f(0), (float, complex))


```