Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Fourier Transform 2021

```python
import numpy as np

def sinc(x):
    return np.sinc(x/np.pi)
```

Note:  if you need to use the sinc function, please use the one provided above rather the one in the NumPy library.

## Part A

Find the Fourier transform for $f(t)$ below.

<img src="fourier_transform2021a.png" width="400"/>

Represent your answer as a function `F(omega)`.

<!-- grade_id: cell-9c81bd6cbb6aa4b7, points: 0, grade: false, solution: true -->
```python
def F(omega):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-019622881b2429a2, points: 2, grade: true, solution: false -->
```python
assert isinstance(F(0), (float, complex))

```

## Part B

Find the inverse transform $f(t)$ if $|F(\omega)|$ and $\angle F(\omega)$ are given below.  

<img src="fourier_transform2021b.png" width="800"/>

Represent your answer as a function `f(t)`.  

<!-- grade_id: cell-3b6f47a4fd843f2b, points: 0, grade: false, solution: true -->
```python
def f(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e8049c8142664a0e, points: 2, grade: true, solution: false -->
```python
assert isinstance(f(0), (float, complex))

```