Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Jacob Ravenberg"
COLLABORATORS = ""
```

---

# Fourier Transform 2022

<!-- grade_id: cell-a4bce8a3c5f78775, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Note:  if you need to use the sinc function, please use the one provided above rather the one in the NumPy library.

## Part A

Find the Fourier transform for $f(t)$ below.

<img src="fourier_transform2022a.png" width="400"/>

Represent your answer as a function `F(omega)`.

<!-- grade_id: cell-5fba9363c411a33f, points: 0, grade: false, solution: true -->
```python
def F(omega):
    if omega == 0:
        return 2
    else:
        return (6*np.sin(omega) - 2*np.sin(omega))/omega
```

<!-- grade_id: cell-94660ab0cea1cc22, points: 2, grade: true, solution: false -->
```python
assert isinstance(F(0), (float, complex))


```

## Part B

Find the inverse transform $f(t)$ if $|F(\omega)|$ and $\angle F(\omega)$ are given below.  

<img src="fourier_transform2022b.png" width="800"/>

Represent your answer as a function `f(t)`.  

<!-- grade_id: cell-94532d69b454180f, points: 0, grade: false, solution: true -->
```python
def f(t):
    if t == -1/2:
        return 3/np.pi
    else:
        return ((np.sin(2*t+1))/(np.pi*(t+(1/2))) + ((1-np.cos(2*t+1))/(2*np.pi*((t+(1/2))**2)))
```

<!-- grade_id: cell-6506e3cf4a2fa1d3, points: 2, grade: true, solution: false -->
```python
assert isinstance(f(0), (float, complex))


```